# Descriptive Stats

For the questions below, we're going to refer exclusively to an HHCAPS survey data set that is available in `/data/hhcaps.csv`.  Use whatever commands you want to calculate the information required to get to the answer.


Put your solution as the last line right before the test cell, using the form shown below for computing the number of rows in a dataframe.

```
answer = my_df.shape[0]
```

The assertions will often give you what the final answer should be, but you won't receive any points unless you compute the answer using code.  For instance, if you just typed `answer = 132` for the numer of rows instead of `answer = my_df.shape[0]`, you would not receive credit.


```python
%matplotlib inline
import pandas as pd
hhcaps = pd.read_csv('/data/hhcaps.csv')
```

## #01 - 

How many columns does this file contain?

```python
hhcaps.shape
```

Output:
```
(12000, 39)
```

<!-- grade_id: q01-solution, points: 0, grade: false, solution: true -->
```python
answer = hhcaps.shape[1]
## Solution goes here
```

<!-- grade_id: q01-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == int)
assert(answer == 39)
```

## #02 - 

How many different values for State are there in this file?

<!-- grade_id: q02-solution, points: 0, grade: false, solution: true -->
```python
answer = hhcaps['State'].nunique()
## Solution goes here
```

<!-- grade_id: q02-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == int)
assert(answer == 55)
```

## #03 - 

Which of those State values has the highest frequency of occurence?

<!-- grade_id: q03-solution, points: 0, grade: false, solution: true -->
```python
answer = hhcaps['State'].value_counts().index[0]
## Solution goes here
```

<!-- grade_id: q03-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == str)
assert(answer == 'TX')
```

## #04 - 

Which of those State values has the best average performance on the `Star Rating for health team communicated well with them` score?

<!-- grade_id: q04-solution, points: 0, grade: false, solution: true -->
```python
answer = hhcaps.groupby(['State']).agg({'Star Rating for health team communicated well with them':'mean'}).reset_index().sort_values('Star Rating for health team communicated well with them', ascending=False).head(1)['State'].values[0]
## Solution goes here
```

<!-- grade_id: q04-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == str)
assert(answer == 'ME')
```

## #05 - 

What was the average score on `Star Rating for how patients rated overall care from agency` for providers listed as having a `Type of Ownership` of `Hospital Based Program`

<!-- grade_id: q05-solution, points: 0, grade: false, solution: true -->
```python
f5 = hhcaps[hhcaps['Type of Ownership']=='Hospital Based Program']
answer = f5['Star Rating for how patients rated overall care from agency'].mean()
```

<!-- grade_id: q05-test, points: 2, grade: true, solution: false -->
```python
import numpy
assert(type(answer) == float or type(answer) == numpy.float64)
assert(round(answer,5) == round(3.6925207756232687,5))
```

## #06 - 

There is one column in the file that has the same value on every row.  What is the name of that column?

<!-- grade_id: q06-solution, points: 0, grade: false, solution: true -->
```python
cols = hhcaps.columns
for col in cols:
    if hhcaps[col].nunique() == 1:
        answer = col    
## Solution goes here
```

<!-- grade_id: q06-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == str)
assert(answer.lower() == 'Offers Nursing Care Services'.lower())
```

## #07 - 

There are six (6) columns in the file that indicate (True or False) if the provider offers certain services.  Those column names all start with `Offers...`.  How many facilities offer **everything** that the survey was interested in asking about?  That is, how many rows have True in all six of those columns?

<!-- grade_id: q07-solution, points: 0, grade: false, solution: true -->
```python
answer = 0
cols = hhcaps.columns
offer_cols=[]
for col in cols:
    if 'Offers' in col:
        offer_cols.append(col)
df_7 = hhcaps[offer_cols]
for index, row in df_7.iterrows():
    counter = 0
    for i in row:
        if i == True:
            counter=counter+1
    if counter == 6:
        answer = answer+1
```

<!-- grade_id: q07-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == int)
assert(answer == 9029)
```

## #08 - 

The survey results report `Number of completed Surveys` as one of the metrics.  What is the `median` of the number of completed surveys per facility?

Note that some rows don't have a valid number in them for the number of completed surveys.


```python
#hhcaps['Number of completed Surveys']=hhcaps['Number of completed Surveys'].replace('Not Available', 0)
df_8 = hhcaps[hhcaps['Number of completed Surveys']!='Not Available']
answer = df_8['Number of completed Surveys'].median()
```

<!-- grade_id: q08-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == float or type(answer) == numpy.float64)
assert(answer == 84)
```

## #09 - 

How many providers are there from St. Louis, Missouri?

*Note that the City columnd may have trailing spaces that need to be accomodated or cleaned up*

```python
df_9 = hhcaps[hhcaps['State'] == 'MO']
df_9_2 = pd.DataFrame(hhcaps['City'].value_counts()).reset_index()
df_9_2['index']=df_9_2['index'].map(lambda x: x.strip())
answer = int(df_9_2[df_9_2['index']== 'SAINT LOUIS']['City'].values[0])
```

<!-- grade_id: q09-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == int)
assert(answer == 23)
```

# #10 - 

Which state got the highest percentage of 'top box' scorings, using the `Percent of patients who gave their home health agency a rating of 9 or 10 on a scale from 0 (lowest) to 10 (highest)` field.

Ignore those rows where the data is **Not Available**

```python
def fun_10(a):
    sum = 0
    for i in a:
        sum = sum + int(i)
    return (sum/len(a))   
```

```python
answer = 0
field = 'Percent of patients who gave their home health agency a rating of 9 or 10 on a scale from 0 (lowest) to 10 (highest)'
df_10 = hhcaps[hhcaps[field] != 'Not Available']
states =df_10['State'].unique()
answer_10 = {}
for state in states:
    df_state = df_10[df_10['State'] == state]
    vals = df_state[field].values
    #top = df_state[field].max()
    #answer_10 [state] = len(df_state[df_state[field] == top])   
    answer_10[state] = fun_10(vals)
answer = sorted(answer_10, key=answer_10.get)[-1]
```

<!-- grade_id: q10-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == str)
assert(answer == 'MS')
```